# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print

```julia
println("I'm excited to learn Julia!")
```

Output:
```
I'm excited to learn Julia!
```

## How to assign variables

All we need is a variable name, value, and an equal's sign!<br>
Julia will figure out types for us.

```julia
my_answer = 42
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

Output:
```
String
```

   To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
🔹="42"
typeof(🔹)
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

```julia
typeof(😺)
```

Note: Julia allows us to write super generic code, and 😺 is an example of this. 

This allows us to write code like

```julia
😀 = 0
😞 = -1
```

```julia
😺 + 😞 
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

```julia
difference = 10 - 3
```

```julia
product = 20 * 5
```

```julia
quotient = 100 / 10
typeof(quotient)
```

Output:
```
Float64
```

```julia
power = 10 ^ 2
```

```julia
modulus = 101 % 2
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
?convert
```

Output:
```
search: [0m[1mc[22m[0m[1mo[22m[0m[1mn[22m[0m[1mv[22m[0m[1me[22m[0m[1mr[22m[0m[1mt[22m [0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m @[0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m

[36m  convert(T, x)[39m

  Convert [36mx[39m to a value of type [36mT[39m.

  If [36mT[39m is an [36mInteger[39m type, an [36mInexactError[39m will be raised if [36mx[39m is not
  representable by [36mT[39m, for example if [36mx[39m is not integer-valued, or is outside
  the range supported by [36mT[39m.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> convert(Int, 3.0)[39m
[36m  3[39m
[36m  [39m
[36m  julia> convert(Int, 3.5)[39m
[36m  ERROR: InexactError: Int64(3.5)[39m
[36m  Stacktrace:[39m
[36m  [...][39m

  If [36mT[39m is a [36mAbstractFloat[39m or [36mRational[39m type, then it will return the closest
  value to [36mx[39m representable by [36mT[39m.

[36m  julia> x = 1
```

```julia
convert(Float64,3.99999999999555554477216342163214680231794)
```

Output:
```
3.9999999999955556
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia
?@assert
```

Output:
```
[36m  @assert cond [text][39m

  Throw an [36mAssertionError[39m if [36mcond[39m is [36mfalse[39m. Preferred syntax for writing
  assertions. Message [36mtext[39m is optionally displayed upon assertion failure.

[33m[1m  │ [22m[39m[33m[1mWarning[22m[39m
[33m[1m  │[22m[39m
[33m[1m  │[22m[39m  An assert might be disabled at various optimization levels. Assert
[33m[1m  │[22m[39m  should therefore only be used as a debugging tool and not used for
[33m[1m  │[22m[39m  authentication verification (e.g., verifying passwords), nor
[33m[1m  │[22m[39m  should side effects needed for the function to work correctly be
[33m[1m  │[22m[39m  used inside of asserts.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> @assert iseven(3) "3 is an odd number!"[39m
[36m  ERROR: AssertionError: 3 is an odd number![39m
[36m  [39m
[36m  julia> @assert isodd(3) "What even are numbers?"[39m
```

```julia
days=365
days_float=convert(Float64,days)
```

Output:
```
365.0
```

```julia
@assert days_float == 365.0
```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
convert(Int64, "1")
```

Output:
```
Stacktrace:
 [1] top-level scope at In[17]:1
```

```julia
parse(Int64, "1")
```

Output:
```
1
```

```julia
?parse
```

Output:
```
search: [0m[1mp[22m[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m[0m[1me[22m try[0m[1mp[22m[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m[0m[1me[22m [0m[1mp[22m[0m[1ma[22m[0m[1mr[22mtial[0m[1ms[22mortp[0m[1me[22mrm [0m[1mp[22m[0m[1ma[22m[0m[1mr[22mtial[0m[1ms[22mortp[0m[1me[22mrm! [0m[1mp[22m[0m[1ma[22mi[0m[1mr[22m[0m[1ms[22m ski[0m[1mp[22mch[0m[1ma[22m[0m[1mr[22m[0m[1ms[22m

[36m  parse(type, str; base)[39m

  Parse a string as a number. For [36mInteger[39m types, a base can be specified (the
  default is 10). For floating-point types, the string is parsed as a decimal
  floating-point number. [36mComplex[39m types are parsed from decimal strings of the
  form [36m"R±Iim"[39m as a [36mComplex(R,I)[39m of the requested type; [36m"i"[39m or [36m"j"[39m can also be
  used instead of [36m"im"[39m, and [36m"R"[39m or [36m"Iim"[39m are also permitted. If the string
  does not contain a valid number, an error is ra
```

Please run the @assert code once you are done with the exercises to validate your answers.