# Data structures

Once we start working with many pieces of data at once, it will be convenient for us to store data in structures like arrays or dictionaries (rather than just relying on variables).<br>

Types of data structures covered:
1. Tuples
2. Dictionaries
3. Arrays

<br>
As an overview, tuples and arrays are both ordered sequences of elements (so we can index into them). Dictionaries and arrays are both mutable.
We'll explain this more below!

## Tuples

We can create a tuple by enclosing an ordered collection of elements in `( )`.

Syntax: <br>
```julia
(item1, item2, ...)```

```julia
myfavoriteanimals = ("penguins", "cats", "sugargliders")
```

Output:
```
("penguins", "cats", "sugargliders")
```

We can index into this tuple,

```julia
myfavoriteanimals[1]
```

Output:
```
"penguins"
```

but since tuples are immutable, we can't update it

```julia
myfavoriteanimals[1] = "otters"
```

Output:
```
Stacktrace:
 [1] top-level scope at In[3]:1
```

## Now in 1.0: NamedTuples

As you might guess, `NamedTuple`s are just like `Tuple`s except that each element additionally has a name! They have a special syntax using `=` inside a tuple:

```julia
(name1 = item1, name2 = item2, ...)
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

Output:
```
(bird = "penguins", mammal = "cats", marsupial = "sugargliders")
```

```julia
myfavoriteanimals = (bird = "penguins", mammal = "cats", bird = "sugargliders")
```

Output:
```
Stacktrace:
 [1] top-level scope at C:\Users\Victor\.julia\packages\IJulia\DrVMH\src\kernel.jl:52
```

Like regular `Tuples`, `NamedTuples` are ordered, so we can retrieve their elements via indexing:

```julia
myfavoriteanimals[end]
```

Output:
```
"sugargliders"
```

They also add the special ability to access values by their name:

```julia
myfavoriteanimals.bird
```

Output:
```
"penguins"
```

## Dictionaries

If we have sets of data related to one another, we may choose to store that data in a dictionary. We can create a dictionary using the `Dict()` function, which we can initialize as an empty dictionary or one storing key, value pairs.

Syntax:
```julia
Dict(key1 => value1, key2 => value2, ...)```

A good example is a contacts list, where we associate names with phone numbers.

```julia
myphonebook = Dict("Jenny" => "867-5309", "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,String} with 2 entries:
  "Jenny"        => "867-5309"
  "Ghostbusters" => "555-2368"
```

In this example, each name and number is a "key" and "value" pair. We can grab Jenny's number (a value) using the associated key

```julia
myphonebook["Jenny"]
```

Output:
```
"867-5309"
```

We can add another entry to this dictionary as follows

```julia
myphonebook["Kramer"] = "555-FILK"
```

Output:
```
"555-FILK"
```

Let's check what our phonebook looks like now...

```julia
myphonebook
```

Output:
```
Dict{String,String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555-FILK"
  "Ghostbusters" => "555-2368"
```

We can delete Kramer from our contact list - and simultaneously grab his number - by using `pop!`

```julia
pop!(myphonebook, "Kramer")
```

Output:
```
"555-FILK"
```

```julia
push!(myphonebook, "Kramer" => "555" )
```

Output:
```
Dict{String,String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555"
  "Ghostbusters" => "555-2368"
```

```julia
myphonebook
```

Output:
```
Dict{String,String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555"
  "Ghostbusters" => "555-2368"
```

Unlike tuples and arrays, dictionaries are not ordered. So, we can't index into them.

```julia
myphonebook[1]
```

Output:
```
Stacktrace:
 [1] getindex(::Dict{String,String}, ::Int64) at .\dict.jl:477
 [2] top-level scope at In[17]:1
```

```julia
myphonebook.Jenny
```

Output:
```
Stacktrace:
 [1] getproperty(::Dict{String,String}, ::Symbol) at .\Base.jl:33
 [2] top-level scope at In[18]:1
```

In the example above, `julia` thinks you're trying to access a value associated with the key `1`.

## Arrays

Unlike tuples, arrays are mutable. Unlike dictionaries, arrays contain ordered collections. <br>
We can create an array by enclosing this collection in `[ ]`.

Syntax: <br>
```julia
[item1, item2, ...]```


For example, we might create an array to keep track of my friends

```julia
myfriends = ["Ted", "Robyn", "Barney", "Lily", "Marshall"]
```

Output:
```
5-element Array{String,1}:
 "Ted"
 "Robyn"
 "Barney"
 "Lily"
 "Marshall"
```

The `1` in `Array{String,1}` means this is a one dimensional vector.  An `Array{String,2}` would be a 2d matrix, etc.  The `String` is the type of each element.

or to store a sequence of numbers

```julia
fibonacci = [1, 1, 2, 3, 5, 8, 13]
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 13
```

```julia
fibonacci[end] = 15
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

```julia
mixture = [1, 1, 2, 3, "Ted", "Robyn"]
```

Output:
```
6-element Array{Any,1}:
 1
 1
 2
 3
  "Ted"
  "Robyn"
```

Once we have an array, we can grab individual pieces of data from inside that array by indexing into the array. For example, if we want the third friend listed in `myfriends`, we write

```julia
myfriends[3]
```

Output:
```
"Barney"
```

We can use indexing to edit an existing element of an array

```julia
myfriends[3] = "Baby Bop"
```

Output:
```
"Baby Bop"
```

Yes, Julia is 1-based indexing, not 0-based like Python.  Wars are fought over lesser issues. I have a friend with the wisdom of Solomon who proposes settling this once and for all with ½ 😃

We can also edit the array by using the `push!` and `pop!` functions. `push!` adds an element to the end of an array and `pop!` removes the last element of an array.

We can add another number to our fibonnaci sequence

```julia
push!(fibonacci, 21)
```

Output:
```
8-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
 21
```

```julia
fibonacci[8] = 21
```

Output:
```
21
```

and then remove it

```julia
pop!(fibonacci)
```

Output:
```
21
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

So far I've given examples of only 1D arrays of scalars, but arrays can have an arbitrary number of dimensions and can also store other arrays. 
<br><br>
For example, the following are arrays of arrays:

```julia
favorites = [["koobideh", "chocolate", "eggs"],["penguins", "cats", "sugargliders"]]
```

Output:
```
2-element Array{Array{String,1},1}:
 ["koobideh", "chocolate", "eggs"]
 ["penguins", "cats", "sugargliders"]
```

```julia
favorites[:][2]
```

Output:
```
3-element Array{String,1}:
 "penguins"
 "cats"
 "sugargliders"
```

```julia
numbers = [[1, 2, 3],[4, 5, 6],[7, 8, 9]]

```

Output:
```
3-element Array{Array{Int64,1},1}:
 [1, 2, 3]
 [4, 5, 6]
 [7, 8, 9]
```

```julia
#println(numbers[2])
println(numbers[:,1])
```

Output:
```
[1, 2, 3, 4, 5, 6, 7, 8, 9]
```

```julia
numbers = [[1, 2, 3] ;[4, 5, 6] ;[7, 8, 9]]
```

Output:
```
9-element Array{Int64,1}:
 1
 2
 3
 4
 5
 6
 7
 8
 9
```

```julia
println(numbers[2])
println(numbers[6])
```

Output:
```
2
6
```

```julia
numbers = [[1 2 3] [4 5 6] [7 8 9]]
```

Output:
```
1×9 Array{Int64,2}:
 1  2  3  4  5  6  7  8  9
```

```julia
println(numbers[2])
println(numbers[2,3])
```

Output:
```
2
Stacktrace:
 [1] getindex(::Array{Int64,2}, ::Int64, ::Int64) at .\array.jl:789
 [2] top-level scope at In[43]:2
```

```julia
numbers = [[1 2 3];[4 5 6];[7 8 9]]
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
println(numbers[2])
println(numbers[2,3])
```

Output:
```
4
6
```

```julia
numbers = [[1 2 3],[4 5 6],[7 8 9]]
```

Output:
```
3-element Array{Array{Int64,2},1}:
 [1 2 3]
 [4 5 6]
 [7 8 9]
```

```julia
println(numbers[2])
println(numbers[2,3])
```

Output:
```
[4 5 6]
Stacktrace:
 [1] getindex(::Array{Array{Int64,2},1}, ::Int64, ::Int64) at .\array.jl:789
 [2] top-level scope at In[47]:2
```

```julia
numbers = [[1, 2, 3] [4, 5, 6] [7, 8, 9]] 
```

Output:
```
3×3 Array{Int64,2}:
 1  4  7
 2  5  8
 3  6  9
```

```julia
println(numbers[2])
println(numbers[2,3])
```

Output:
```
2
8
```

The simplest way to define a matrix

```julia
numbers = [1 2 3; 4 5 6; 7 8 9] 
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 4  5  6
 7  8  9
```

The following will not work as the "," is treated as seperator of array elements and semicolon is not permitted in between unless one of the above tricks of matrix creation is applied.

```julia
numbers = [1, 2, 3; 4, 5, 6; 7, 8, 9] 
```

Output:
```
Stacktrace:
 [1] top-level scope at C:\Users\Victor\.julia\packages\IJulia\DrVMH\src\kernel.jl:52
```

Below are examples of 2D and 3D arrays populated with random values.

```julia
y = rand(4,3)
```

Output:
```
4×3 Array{Float64,2}:
 0.502644  0.180816  0.605967
 0.241358  0.133781  0.492097
 0.860804  0.243546  0.80082
 0.14517   0.303648  0.686514
```

```julia
y[:]
```

Output:
```
12-element Array{Float64,1}:
 0.5026436466874751
 0.24135824527933614
 0.8608035426484164
 0.1451702892508966
 0.18081623808770053
 0.1337812520917394
 0.2435464691443654
 0.3036482233896971
 0.6059668053199894
 0.49209679398228845
 0.8008200127872795
 0.6865140836104189
```

```julia
x = rand(4, 3, 2)
```

Output:
```
4×3×2 Array{Float64,3}:
[:, :, 1] =
 0.423937  0.150621  0.0284276
 0.94715   0.461231  0.550478
 0.690565  0.684341  0.44705
 0.752426  0.625381  0.29683

[:, :, 2] =
 0.53064   0.0198399  0.372656
 0.409091  0.835435   0.417472
 0.942018  0.692341   0.745662
 0.447406  0.178595   0.754698
```

```julia
x[:]
```

Output:
```
24-element Array{Float64,1}:
 0.42393686786689844
 0.947149983540601
 0.6905654618265784
 0.7524262902436758
 0.15062053371336948
 0.46123115076010257
 0.6843409047955222
 0.625380983233002
 0.028427570327122398
 0.5504778973735025
 0.44704977423190373
 0.2968295331962947
 0.5306404322035885
 0.40909055697814156
 0.9420175163778002
 0.44740569852441925
 0.01983991214378955
 0.8354351475350112
 0.6923414499648872
 0.17859513883667377
 0.3726561687142662
 0.417471654491548
 0.7456616786850643
 0.7546978062057383
```

Be careful when you want to copy arrays!

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

```julia
somenumbers = fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

```julia
somenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
 404
   1
   2
   3
   5
   8
  15
```

Editing `somenumbers` caused `fibonacci` to get updated as well!

In the above example, we didn't actually make a copy of `fibonacci`. We just created a new way to access the entries in the array bound to `fibonacci`.

If we'd like to make a copy of the array bound to `fibonacci`, we can use the `copy` function.

```julia
# First, restore fibonacci
fibonacci[1] = 1
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

```julia
somemorenumbers = copy(fibonacci[1:4])
```

Output:
```
4-element Array{Int64,1}:
 1
 1
 2
 3
```

somenumbers = fibonacci[:] also behaves as copy

```julia
somemorenumbers[1] = 404
```

Output:
```
404
```

```julia
fibonacci
```

Output:
```
7-element Array{Int64,1}:
  1
  1
  2
  3
  5
  8
 15
```

In this last example, fibonacci was not updated. Therefore we see that the arrays bound to `somemorenumbers` and `fibonacci` are distinct.

#### Summary of Data Structures

|Data structure   | mutable | ordered |
|---------|:-------:|:-------:|
|Tuples     |   no |    yes    |
|Dictionary |    yes    |    no     |
|Array      |    yes    |    yes    |

### Exercises

#### 3.1 
Create an array, `a_ray`, with the following code:

```julia
a_ray = [1, 2, 3]
```

Add the number `4` to the end of this array and then remove it.

```julia
a_ray = [1, 2, 3]
```

Output:
```
3-element Array{Int64,1}:
 1
 2
 3
```

```julia
push!(a_ray,4)
```

Output:
```
4-element Array{Int64,1}:
 1
 2
 3
 4
```

```julia
pop!(a_ray)
```

Output:
```
4
```

<!-- grade_id: cell-a6b20b78fbfeb449, points: 1, grade: true, solution: false -->
```julia
@assert a_ray == [1, 2, 3]
```

#### 3.2 
Try to add "Emergency" as key to `myphonebook` with the value `string(911)` with the following code
```julia
myphonebook["Emergency"] = 911
```

Why doesn't this work?

```julia
myphonebook

```

Output:
```
Dict{String,String} with 3 entries:
  "Jenny"        => "867-5309"
  "Kramer"       => "555"
  "Ghostbusters" => "555-2368"
```

```julia
typeof(myphonebook)
```

Output:
```
Dict{String,String}
```

```julia
myphonebook["Emergency"] = 911
```

Output:
```
Stacktrace:
 [1] setindex!(::Dict{String,String}, ::Int64, ::String) at .\dict.jl:380
 [2] top-level scope at In[77]:1
```

```julia
myphonebook
```

Output:
```
Dict{String,String} with 4 entries:
  "Jenny"        => "867-5309"
  "Emergency"    => "911"
  "Kramer"       => "555"
  "Ghostbusters" => "555-2368"
```

#### 3.3 
Create a new dictionary called `flexible_phonebook` that has Jenny's number stored as an integer and Ghostbusters' number stored as a string with the following code

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

```julia
flexible_phonebook = Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

```julia
typeof(flexible_phonebook)
```

Output:
```
Dict{String,Any}
```

<!-- grade_id: cell-66f1d636c676c2f2, points: 1, grade: true, solution: false -->
```julia
@assert flexible_phonebook == Dict("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

#### 3.4 
Add the key "Emergency" with the value `911` (an integer) to `flexible_phonebook`.

```julia
flexible_phonebook["Emergency"]=911
```

Output:
```
911
```

```julia
flexible_phonebook
```

Output:
```
Dict{String,Any} with 3 entries:
  "Jenny"        => 8675309
  "Emergency"    => 911
  "Ghostbusters" => "555-2368"
```

<!-- grade_id: cell-ee862ab938680259, points: 1, grade: true, solution: false -->
```julia
@assert haskey(flexible_phonebook, "Emergency")
```

<!-- grade_id: cell-c8c40ababe7cdfce, points: 2, grade: true, solution: false -->
```julia
@assert flexible_phonebook["Emergency"] == 911
```

#### 3.5 
Why can we add an integer as a value to `flexible_phonebook` but not `myphonebook`? How could we have initialized `myphonebook` so that it would accept integers as values?

```julia
myphonebookv2=Dict{String,Any}("Jenny" => 8675309, "Ghostbusters" => "555-2368")
```

Output:
```
Dict{String,Any} with 2 entries:
  "Jenny"        => 8675309
  "Ghostbusters" => "555-2368"
```

```julia
myphonebookv2["Emergency"] = 911
```

Output:
```
911
```

```julia
myphonebookv2
```

Output:
```
Dict{String,Any} with 3 entries:
  "Jenny"        => 8675309
  "Emergency"    => 911
  "Ghostbusters" => "555-2368"
```

Please run @assert cells upon completion of the excercise to validate your answers